## Implementing Back-propagation Algorithm with XOR data

### XOR data: <br>
**$ x_0 \ x_1 \ y$** <br>
$0 \ \ \  0 \ \ \  0$ <br>
$0 \ \ \  1 \ \ \  1$ <br>
$1 \ \ \  0 \ \ \  1$<br>
$1 \ \ \  1 \ \ \  0$<br>

---



```python
import numpy as np
import matplotlib.pyplot as plt
```

##Activation function

Sigmoid function $$\frac{1}{1+ e^{-x}} $$

<!-- grade_id: cell-d84ad9dbcb889c3f, points: 0, grade: false, solution: true -->
```python
#Define our activation function

def sigmoid (x):
    '''
    Input:
        x: numpy array of any shape
    Output:
        y: numpy array of same shape as x
    '''
    # YOUR CODE HERE
    return 1/(1+np.exp(-x))
```

<!-- grade_id: cell-80b53c7b5034f924, points: 1, grade: true, solution: false -->
```python
'''Testing'''
assert sigmoid(0)==0.5
assert np.isclose(sigmoid(-2), 0.119202922, atol=0.0001)
print('Test passed', '\U0001F44D')
```

Output:
```
Test passed 👍
```

<!-- grade_id: cell-9ebb909521c85d2b, points: 0, grade: false, solution: true -->
```python
# Define the activation function derivative

def sigmoid_derivative(x):
    '''
    Input:
        x: numpy array of any shape
    Output:
        y: numpy array of same shape as x
        y = derivative of sigmoid
    '''
    # YOUR CODE HERE
    return x*(1-x)
```

<!-- grade_id: cell-8668ae928d66bf7c, points: 1, grade: true, solution: false -->
```python
'''Testing code for sigmoid_derivative'''
assert sigmoid_derivative(1) == 0
assert sigmoid_derivative(0) == 0
print('Test passed', '\U0001F44D')

```

Output:
```
Test passed 👍
```

## Defining the model

<!-- grade_id: cell-4a7dcd60006d48dc, points: 0, grade: false, solution: true -->
```python
#Define the NeuralNetwork class

class NeuralNetwork:
    def __init__(self, net_arch):
        '''   
        Input:
            net_arch: list of 3 integers
        Action:
            Creates instance variables:
                self.input: np array of shape (ni,1)
                self.layer1: nprarray of shape (nh,1)
                self.output: np array of shape (no,1)
                self.weights1: np array of shape (nh, ni), initialized randomly between (-1,1)
                self.weights2: np array of shape (no, nh), initialized randomly between (-1,1)
                
            NOTE: We do not use bias explicitly here. Input x can have the first element 1 to have a bias term.
        '''
        ni = net_arch[0]  ## Number of neurons in input layer    
        nh = net_arch[1]  ## Number of neurons in hidden layer
        no = net_arch[2]  ## Number of neurons in output layer
        
        self.ni = ni
        self.nh = nh
        self.no = no
        
        # YOUR CODE HERE
        self.input=np.empty((ni,1))
        self.layer1=np.empty((nh,1))
        self.output=np.empty((no,1))
        self.weights1=(2*np.random.rand(nh,ni)-1)/2
        self.weights2=(2*np.random.rand(no,nh)-1)/2
        
    def feedforward(self,x):
        '''
        Input:
            x: numpy array of shape (ni,1)
        Action:
            
        Return:
            output: numpy array of shape (no,1),
        '''
        # YOUR CODE HERE
        h1=np.array([])
        v1=np.array([])
        h2=np.array([])
        h1=self.weights1.dot(x[:(self.ni),:])
        v1=sigmoid(h1)
        self.layer1=v1
        h2=self.weights2.dot(v1)
        output=sigmoid(h2)
        return output
             

    def backprop(self,x,y,eta):
        '''
        # application of the chain rule to find derivative of the loss function with respect to weights2 and weights1
        Input:
            x: numpy array of shape (ni,1)
            y: numpy array of shape (no,1)
            eta: learning rate
        Action:
        # Finding the derivatives
            del_weights2: np array of shape (no,nh) that stores the derivative of the loss function with respect to weights2
            del_weights1: np array of shape (nh,ni) that stores the derivative of the loss function with respect to weights1
            
        # Update the weights with the derivative of the loss function
            weights1 += eta*del_weights1
            weights2 += eta*del_weights2
        '''
        e=y-self.output
        del_weights2=np.dot((e*sigmoid_derivative(y)),self.layer1.T)
        print(del_weights2.shape)
        print(self.layer1)
        print(x)
        del_weights1=np.dot(((sigmoid_derivative(self.layer1.T))*del_weights2*self.weights2).T,x.T)
        self.weights2+=eta*del_weights2
        self.weights1+=eta*del_weights1
   


    def fit(self, X, Y, eta, epochs):
        '''
        input:
        X: training input data of shape (4,2)
        Y: training output of shape (4,1)
        eta: learning rate
        epochs: number of epochs
        Action:
        # Modify the input by adding ones of shape(4,1) 
        # Set up the feed-forward propagation for the modified input   
        # Set up the back-propagation of the error to adjust the weights
        '''
        # YOUR CODE HERE
        for i in range(epochs):
          #X=np.hstack((X,(np.ones(4,1))))
          a = np.ones((4,1))
          x = np.add(X,a)
          self.output=feedforward(self,x)
          backprop(self,X,Y,eta)
          
          
        
        
        
    def predict(self,x,y):
        '''
        # Predict function is used to check the prediction result of the neural network
        Input:
        x: single input data of shape (1,3)
        y: single output data of shape (1,1)
        Action
        pred_out: predict the output based on the model using feedforward
        
        Output
        error: y - pred_out
        
        
        '''
        # YOUR CODE HERE
        y1=feedforward(self,x.T)
        error=y1-y
        return(error)
        
```

<!-- grade_id: cell-640694d6a41609e8, points: 2, grade: true, solution: false -->
```python
'''Testing code for __init__'''

net_arch = [3,4,1]
nn1 = NeuralNetwork(net_arch)
assert nn1.input.shape==(3,1)
assert nn1.layer1.shape == (4,1)
assert nn1.output.shape == (1,1)
assert np.all(nn1.weights1 < 1)
print('Test passed', '\U0001F44D')

```

Output:
```
Test passed 👍
```

<!-- grade_id: cell-f0a271c06ea8c25b, points: 2, grade: true, solution: false -->
```python
'''Testing code for feedforward'''

def feedforward_original(nn1,x):
    assert x.shape == (nn1.ni, 1)
    layer1 = sigmoid(np.dot(nn1.weights1, x))
    output = sigmoid(np.dot(nn1.weights2, layer1))
    return output
x = np.array([0,1,1]).reshape(-1, 1)
assert nn1.feedforward(x) == feedforward_original(nn1, x)
print('Test passed', '\U0001F44D')
```

Output:
```
Test passed 👍
```

<!-- grade_id: cell-a4a2893378d844f3, points: 4, grade: true, solution: false -->
```python
'''Testing code for backprop'''
def backprop_original(nn1,x,y,eta):
    weights1 = nn1.weights1
    weights2 = nn1.weights2
    del_weights2 = np.dot(((y - nn1.output) * sigmoid_derivative(nn1.output)),nn1.layer1.reshape(-1, 1).T)
    del_weights1 = np.dot(((y - nn1.output) * sigmoid_derivative(nn1.output)*nn1.weights2.T * sigmoid_derivative(nn1.layer1)), x.T)

    # update the weights with the derivative (slope) of the loss function
    weights1 += eta*del_weights1
    weights2 += eta*del_weights2
    return(weights1, weights2)

x = np.array([0,1,1]).reshape(-1, 1)
y = np.array([[0],])
eta = 1
nn1.backprop(x, y, eta)
w1, w2 = backprop_original(nn1, x, y, eta) 
assert np.all(np.isclose(w1, nn1.weights1))
assert np.all(np.isclose(w2, nn1.weights2))
print('Test passed', '\U0001F44D')
```

Output:
```
(1, 4)
[[0.31316231]
 [0.53597592]
 [0.51656296]
 [0.41042472]]
[[0]
 [1]
 [1]]
Test passed 👍
```

## Fitting the data (Training)

<!-- grade_id: cell-e59bb4a5a7ddab07, points: 0, grade: false, solution: true -->
```python
## CHECK THE PERFORMANCE
'''
Input:
# Set the input data
X = ([[0.1, 0.1], [0.1, 0.9],
                [0.9, 0.1], [0.9, 0.9]])
# Set the labels, the correct results for the xor operation
Y = ([[0.1], [0.9], 
                 [0.9], [0.1]])
Action:
# Initialize the NeuralNetwork with
# 3 input neurons
# 4 hidden neurons
# 1 output neuron

# Fit the datas
'''
# YOUR CODE HERE
X = [[0.1, 0.1], [0.1, 0.9],[0.9, 0.1], [0.9, 0.9]]
x = np.array(X)
Y = ([[0.1], [0.9],[0.9], [0.1]])
y = np.array(Y)
arch = np.array([3,4,1])
neural = NeuralNetwork(arch)
epochs =1000
eta =0.06
neural.fit(x, y, eta,epochs)

```

Output:
```
<ipython-input-67-cc45126e2c5e> in <module>()
     24 epochs =1000
     25 eta =0.06
---> 26 neural.fit(x, y, eta,epochs)

<ipython-input-63-26f4ec275ba2> in fit(self, X, Y, eta, epochs)
     95           a = np.ones((4,1))
     96           x = np.add(X,a)
---> 97           self.output=feedforward(self,x)
     98           backprop(self,X,Y,eta)
     99 

NameError: name 'feedforward' is not defined
```

<!-- grade_id: cell-7391ad1a93e30118, points: 3, grade: true, solution: false -->
```python
'''Testing for fit'''
X = np.array([[0.1, 0.1], [0.1, 0.9],
                [0.9, 0.1], [0.9, 0.9]])
# Set the labels, the correct results for the xor operation
Y = np.array([[0.1], [0.9], 
                 [0.9], [0.1]])
nn1.fit(X,Y,1,10000)
x = np.array([1,1,1]).reshape(-1, 1)
y = np.array([[0],])
print(nn1.feedforward(x),y)
assert np.all(np.isclose(nn1.feedforward(x),y,atol=0.1))
print('Test passed', '\U0001F44D')
```

## Plotting 

```python
def plotting(X, Y):
  x_plot = X.T
  color = []
  for i in Y:
    if i[0] == 0:
      color.append('g')
    else:
      color.append('r')
  color = np.array(color)
  print(x_plot)
  plt.figure()
  plt.xlabel('x1')
  plt.ylabel('x2')
  plt.scatter(x_plot[0],x_plot[1],color=color)
  plt.show()
```

```python
plotting(X, Y)
```

## Could you test it now?

Find the error between the predicted output and the desired output.

```python
def testing(X, Y):
  ones = 0.9*np.ones((X.shape[0],1))
  x_test = np.concatenate([ones, X], axis=1)
  y_test = Y

  for k in range(4):
    print(nn.predict(x_test[k].reshape(-1, 1),y_test[k]))
    
```

<!-- grade_id: cell-08caa0ed7b0ce465, points: 2, grade: true, solution: false -->
```python
'''Testing the prediction'''
x = np.array([0.9,0.9,0.9]).reshape(-1, 1)
y = np.array([[0.1],])

assert np.all(np.isclose(nn1.predict(x,y),0, atol=0.01))
print('Test passed', '\U0001F44D')
```

# Advanced
## Does the performance increase with increasing the number of neurons in the hidden layer?
- Repeat the training with 1 neuron in the hidden layer, then with 3 neuron and then with 5 neuron in the hidden layer to see the trend in performance
- Compare the training error
- Compare the testing error

```python

```